# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # YOUR CODE HERE
    res = True
    x_list = int_to_list(x)
    y_list = int_to_list(y)
    
    res = check_digits(x_list, y)
    if (res != False):
        res = check_digits(y_list, x)
    return res
```

```python
def int_to_list(number):
    string = str(number)
    str_list = []
    
    for char in string:
        str_list.append(char)
    return str_list
```

```python
def check_digits(char_list, string):
    res = True
    
    for char in char_list:
        if (string.find(char) == -1):
            res = False
    return res
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
x = 1

while(same_digits(str(2*x), str(3*x)) == False or same_digits(str(3*x), str(4*x)) == False or same_digits(str(5*x), str(6*x)) == False or same_digits(str(6*x), str(2*x)) == False):
    x += 1
    
print(x)
    
```

Output:
```
142857
```

```python

```