# MLE for Poisson Distribution

In this notebook you will perform [maximum likelihood estimation](https://en.wikipedia.org/wiki/Maximum_likelihood), or MLE, for the parameter $\lambda$ in the [Poisson](https://en.wikipedia.org/wiki/Poisson_distribution) distribution and bootstrap the results.

## Imports

```python
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
from scipy.optimize import minimize
from scipy.stats import poisson
```

```python
import tensorflow as tf
from tensorflow.contrib import distributions as dist
```

## Create a dataset

Generate a dataset from the Poisson distribution with a known parameter $\lambda$:

```python
np.random.seed(0)
位 = 3.0
data = np.random.poisson(位, size=20)
data
```

Output:
```
array([5, 6, 2, 0, 7, 3, 3, 1, 6, 1, 3, 3, 2, 2, 2, 2, 3, 3, 3, 1])
```

## MLE with Tensorflow

Use TensorFlow to write a function that performs MLE for the $\lambda$ parameter of the Poisson distribution. In your function you should:

* Create all of the needed TensorFlow `Variable` and `placeholder` instances.
* Create an initialize a `Session`.
* Call `Session.close()` before returning your estimate for lambda (you can also create and use the `Session` with Python's `with` statement).
* Use the `float32` dtype.
* Play with the learning parameter and number of gradient descent steps to get the function to run quickly and give results that are accurate enough to pass the tests.

```python
dist.Poisson??
```

<!-- grade_id: mle-poisson-b, points: 0, grade: false, solution: true -->
```python
def fit_lambda(data):
    """Perform MLE to estimate the 位 parameter of the Poisson distribution.
    
    Parameters
    ----------
    data: ndarray
        The data to use in estimating lambda
    
    Returns
    -------
    lambda: float
        The MLE value for lambda
    """
    # YOUR CODE HERE
    lam = tf.Variable(1.0, dtype=tf.float64)
    x = tf.placeholder(dtype=tf.float64)
    model = dist.Poisson(lam)
    nll = tf.reduce_sum(-1.0*model.log_pdf(x))
    sess = tf.Session()
    init = tf.global_variables_initializer()
    sess.run(init)
    
```

<!-- grade_id: mle-poisson-c, points: 4, grade: true, solution: false -->
```python
assert abs(fit_lambda(3*np.ones(10))-3.0) < 1e-3
assert abs(fit_lambda(100*np.ones(100))-100.0) < 1e-3
assert abs(fit_lambda(data)-data.mean()) < 1e-3
```

Output:
```
<ipython-input-14-00ed375a9800> in fit_lambda(data)
     16     x = tf.placeholder(dtype=tf.float64)
     17     model = dist.Poisson(lam)
---> 18     nll = tf.reduce_sum(-1.0*model.log_pdf(x))
     19     sess = tf.Session()
     20     init = tf.global_variables_initializer()

/opt/conda/lib/python3.5/site-packages/tensorflow/contrib/distributions/python/ops/distribution.py in log_pdf(self, value, name, **condition_kwargs)
    829                   PendingDeprecationWarning)
    830     if not self.is_continuous:
--> 831       raise TypeError("log_pdf is undefined for non-continuous distributions.")
    832     return self.log_prob(value, name=name, **condition_kwargs)
    833 

TypeError: log_pdf is undefined for non-continuous distributions.
```

Here is the MLE for lambda given our original data:

```python
fit_lambda(data)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-10-c588cd1c0aa7> in <module>()
----> 1 fit_lambda(data)

NameError: name 'fit_lambda' is not defined
```

And the time for a single fit:

```python
%timeit -n1 -r1 fit_lambda(data)
```

## MLE with `scipy.optimize.minimize`

Use SciPy, write a function that performs MLE for the $\lambda$ parameter of the Poisson distribution. In your function you should:

* Use `scipy.stats.poisson` to calculate the negative log-likelihood of the Poisson distribution.
* Use `scipy.optimize.minimize` to minimize the negative log-likelihood.

<!-- grade_id: mle-poisson-f, points: 0, grade: false, solution: true -->
```python
def fit_lambda_fast(data):
    """Perform MLE to estimate the 位 parameter of the Poisson distribution.
    
    Parameters
    ----------
    data: ndarray
        The data to use in estimating lambda
    
    Returns
    -------
    lambda: float
        The MLE value for lambda
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: mle-poisson-g, points: 2, grade: true, solution: false -->
```python
assert abs(fit_lambda_fast(3*np.ones(10))-3.0) < 1e-3
assert abs(fit_lambda_fast(100*np.ones(100))-100.0) < 1e-3
assert abs(fit_lambda_fast(data)-data.mean()) < 1e-3
```

Here is the estimated value of $\lambda$:

```python
fit_lambda_fast(data)
```

And the timing of a single run, which should be much faster than the TensorFlow result:

```python
%timeit -n1 -r1 fit_lambda_fast(data)
```

## Bootstrap

We see that the MLE for $\hat{\lambda}$ is close to the original value of $\lambda=3.0$. Bootstrap resample this estimator 200 times (using your faster version) to find the distribution of $\hat{\lambda}$. Save the distributions of $\lambda$ values in an list or array named `estimates`:

<!-- grade_id: mle-poisson-j, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Plot the distibution of bootstrapped estimates using Matplotlib. Lable your axes, use a grid and customize the bins:

<!-- grade_id: mle-poisson-k, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Compute the print the 95% confidence interval:

<!-- grade_id: mle-poisson-l, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```