# Basic Vector Operations

In this notebook you will implement a number of basic vector operations using pure Python.

```python
import math
```

Implement a function that computes the sum of two input vectors `x` and `y`:


<!-- grade_id: basicvectora, points: 0, grade: false, solution: true -->
```python
def vector_sum(x, y):
    """Compute the vector sum of x and y as a Python list."""
    # YOUR CODE HERE
    res = []
    for index in range(len(x)):
        res.append(x[index] + y[index])
        
    return res
```

<!-- grade_id: basicvectorb, points: 1, grade: true, solution: false -->
```python
assert vector_sum([0,0,0,0], [2,2,2,2])==[2,2,2,2]
assert vector_sum([1,2], [-1,-2])==[0,0]
```

Implement a function that computes the difference of two input vectors `x` and `y`:

<!-- grade_id: basicvectorc, points: 0, grade: false, solution: true -->
```python
def vector_diff(x, y):
    """Compute the vector difference of x and y as a Python list."""
    # YOUR CODE HERE
    res = []
    for index in range(len(x)):
        res.append(x[index] - y[index])
        
    return res
```

<!-- grade_id: basicvectord, points: 1, grade: true, solution: false -->
```python
assert vector_diff([1,1], [0,0])==[1,1]
assert vector_diff([1,2,3,4],[1,2,3,4])==[0,0,0,0]
```

Implement a function that computes the product of a scalar `a` with a vectors `x`:

<!-- grade_id: basicvectore, points: 0, grade: false, solution: true -->
```python
def vector_scalar_mul(a, x):
    """Compute a*x as a Python list."""
    # YOUR CODE HERE
    res = []
    for index in range(len(x)):
        res.append(x[index] * a)
    
    return res
```

<!-- grade_id: basicvectorf, points: 1, grade: true, solution: false -->
```python
assert vector_scalar_mul(1.0, [1,2,3,4])==[1,2,3,4]
assert vector_scalar_mul(2.0, [-1,0,1])==[-2,0,2]
```

Implement a function that computes the dot product between two vectors `x` and `y`:

<!-- grade_id: basicvectorg, points: 0, grade: false, solution: true -->
```python
def vector_dot(x, y):
    """Compute the dot product of vectors a and b."""
    # YOUR CODE HERE
    res = 0
    for index in range(len(x)):
        res += (x[index] * y[index])
    return res
```

<!-- grade_id: basicvectorh, points: 1, grade: true, solution: false -->
```python
assert vector_dot([0,1],[1,0])==0.0
assert vector_dot([1,0],[1,0])==1.0
assert vector_dot([1,1],[-1,1])==0.0
```

Implement a function that computes the $L_1$ norm of a vector `x`:

<!-- grade_id: basicvectori, points: 0, grade: false, solution: true -->
```python
def L1(x):
    """Compute the L1 norm of x."""
    # YOUR CODE HERE
    res = 0
    for index in x:
        res += math.fabs(index)
    return res
```

<!-- grade_id: basicvectorj, points: 1, grade: true, solution: false -->
```python
assert L1([0,0])==0
assert L1([1,1])==2
```

Implement a function that computes the $L_2$ norm of a vector `x`:

<!-- grade_id: basicvectork, points: 0, grade: false, solution: true -->
```python
def L2(x):
    """Compute the L2 norm of x."""
    # YOUR CODE HERE
    res = 0
    for index in x:
        res += math.pow(math.fabs(index),2)
    return math.sqrt(res)
```

<!-- grade_id: basicvectorl, points: 1, grade: true, solution: false -->
```python
assert L2([0,0])==0
assert L2([1,1])==math.sqrt(2)
```

Implement a function that computes the Mean Squared Error (MSE) between two vectors `x` and `y`:

<!-- grade_id: basicvectorm, points: 0, grade: false, solution: true -->
```python
def mse(x, y):
    """Compute the mean squared error between x and y."""
    # YOUR CODE HERE
    res = 0
    count = 0
    for index in range(len(x)):
        res += math.pow(math.fabs(x[index] - y[index]), 2)
        count += 1
    return res / count
```

<!-- grade_id: basicvectorn, points: 1, grade: true, solution: false -->
```python
assert mse([0,0],[1,1])==1.0
assert mse([1,2,3,4],[5,6,7,8])==16.0
```

Implement a function that computes the mean of all the elements of a vector `x`:

<!-- grade_id: basicvectoro, points: 0, grade: false, solution: true -->
```python
def mean(x):
    """Compute the mean of the elements of x."""
    # YOUR CODE HERE
    res = 0
    count = 0
    for index in x:
        res += index
        count += 1;
    return res/count
```

<!-- grade_id: basicvectorp, points: 1, grade: true, solution: false -->
```python
assert mean([1,1,1,1])==1
assert mean([1,2,3,4])==2.5
```

Implement a function that computes the variance of the elements of a vector `x`:

<!-- grade_id: basicvectorq, points: 0, grade: false, solution: true -->
```python
def var(x):
    """Compute the variance of the elements of x."""
    # YOUR CODE HERE
    first_mean = 0
    for num in x:
        first_mean += num
    first_mean = first_mean / len(x)
    end_mean = 0
    for num in x:
        end_mean += pow(num - first_mean, 2)

    return end_mean / len(x)
```

<!-- grade_id: basicvectorr, points: 1, grade: true, solution: false -->
```python
assert var([1,1,1,1])==0.0
assert var([1,2,3,4])==1.25
```

Implement a function that computes the standard deviation of a vector `x`:

<!-- grade_id: basicvectors, points: 0, grade: false, solution: true -->
```python
def std(x):
    """Compute the standard deviation of the elements of x."""
    # YOUR CODE HERE
    first_mean = 0
    for num in x:
        first_mean += num
    first_mean = first_mean / len(x)
    end_mean = 0
    for num in x:
        end_mean += pow(num - first_mean, 2)

    return math.sqrt(end_mean / len(x))
```

<!-- grade_id: basicvectort, points: 1, grade: true, solution: false -->
```python
assert abs(std([2,4,6,8])-2.23606)<0.0001
```

```python

```