# Pandas Missing Values

## Imports and data

```python
import numpy as np
import pandas as pd
```

```python
df = pd.DataFrame(np.ones((10,3)), index=list('abcdefghij'), columns=['one', 'two', 'three'])
df.loc['a','one':'two'] = np.nan
df.loc['c','one'] = np.nan
df.loc['d','two'] = np.nan
df.loc['e','three'] = np.nan
df.loc['f',:] = np.nan
df.loc['g','one':'two'] = np.nan
df.loc['h', 'two':'three'] = np.nan
df['state'] = ['CA', '', None, 'OR', 'WA', None, '', 'WA', 'OR', None]
```

Here is the `DataFrame` we will work with:

```python
df
```

Replace empty strings in the `state` column by `None`, so that its missing values are handled in a consistent manner:

<!-- grade_id: missing-values-c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: missing-values-d, points: 2, grade: true, solution: false -->
```python
assert '' not in df.state.unique()
```

Create a new `DataFrame`, named `df2`, that has all rows with any missing values dropped:

<!-- grade_id: missing-values-e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: missing-values-f, points: 2, grade: true, solution: false -->
```python
assert len(df2)==2
assert 'b' in df2.index
assert 'i' in df2.index
```

Create a new `DataFrame`, named `df3`, from `df1` by dropping rows that have only missing values:

<!-- grade_id: missing-values-g, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: missing-values-h, points: 2, grade: true, solution: false -->
```python
assert len(df3)==9
```

Create a new `DataFrame`, named `df4`, from `df1` that has all columns with fewer than 7 actual values dropped:

<!-- grade_id: missing-values-i, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: missing-values-j, points: 2, grade: true, solution: false -->
```python
assert list(df4.columns)==['three']
```

Create a new `DataFrame`, named `df5`, from `df1` that has only the numerical columns, with missing values replace by the number -9.

<!-- grade_id: missing-values-k, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: missing-values-l, points: 2, grade: true, solution: false -->
```python
assert list(df5.columns)==['one','two','three']
sums = df5.sum()
assert sums['one']==-30.0
assert sums['two']==-40.0
assert sums['three']==-20.0
```

Write a function `count_null` that takes a `Series` and return an integer valued count of the number of null values in the `Series`:

<!-- grade_id: missing-values-m, points: 0, grade: false, solution: true -->
```python
def count_null(column):
    """Count the number of missing values in a column (Series)."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: missing-values-n, points: 2, grade: true, solution: false -->
```python
assert count_null(df.one)==4
assert count_null(df.two)==5
assert count_null(df.three)==3
assert count_null(df.state)==4
```